```python
# Initialize Otter
import otter
grader = otter.Notebook("hw08.ipynb")
```

<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Homework 8: A/B Testing

## References

* [Comparing Two Samples](https://inferentialthinking.com/chapters/12/Comparing_Two_Samples.html)
* [`datascience` Documentation](https://datascience.readthedocs.io/)
* [Python Quick Reference](https://ccsf-math-108.github.io/materials-fa23/resources/quick_reference.html)

---

## Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your Tasks are categorized as auto-graded (📍) and manually graded (📍🔎).
    - For all the auto-graded tasks:
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to run some tests on your code.
        - Keep in mind that for homework and project assignments, sometimes there are hidden tests that you will not be able to see the results of that we use for scoring the correctness of your response. **Passing the auto-grader does not guarantee that your answer is correct.**
    - For all the manually graded tasks:
        - You might need to provide your own response to the provided prompt. Do so by replacing the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or something else using code. Do so by replacing the `...` in the code cell to generate the image, table, etc.
        - In either case, review the rubric on the associated <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page to understand the scoring criteria.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

---

## Configure the Notebook

Run the following cell to configure this Notebook.

```python
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

---

## Who is Older?

---

In an attempt to answer the question 'Who is Older?', data scientists have drawn a simple random sample of size 500 from a large population of adults. Each member of the population happened to identify as either "male" or "female". Data was collected on several attributes of the sampled people, including age. The table `sampled_ages` contains one row for each person in the sample, with columns containing the individual's gender identity and their age.

```python
sampled_ages = Table.read_table('age.csv')
sampled_ages.show(5)
```

---

### Task 01 📍

How many females are there in this sample? Please use the provided skeleton code.

*Hint:* Keep in mind that `.group` sorts categories in alphabetical order!

_Points:_ 3

```python
num_females = sampled_ages.group(...)...
num_females
```

```python
grader.check("task_01")
```

---

### Task 02 📍

Complete the cell below so that `avg_male_vs_female` evaluates to `True` if the sampled males are older than the sampled females on average, and `False` otherwise. Use Python code to achieve this.

_Points:_ 3

```python
group_mean_tbl = sampled_ages.group(...)
group_means = group_mean_tbl...       # array of mean ages
avg_male_vs_female = group_means... > group_means...
avg_male_vs_female
```

```python
grader.check("task_02")
```

---

### Task 03 📍

The data scientists want to use the data to test whether males are older than females on average or whether the ages of the two groups have the same distribution. 

In the code cell below, assign `null_statement_number` and `alternative_statement_number` to the number corresponding to the correct statement from the list below. (There is only one correct answer for each.)

1. In the sample, the males and females have the same distribution of ages; the sample averages of the two groups are different due to chance.
2. In the population, the males and females have the same distribution of ages; the sample averages of the two groups are different due to chance.
3. The age distributions of males and females in the population are different due to chance.
4. The males in the sample are older than the females, on average.
5. The males in the population are older than the females, on average.
6. The average ages of the males and females in the population are different.

_Points:_ 3

```python
null_statement_number = ...
alternative_statement_number = ...
```

```python
grader.check("task_03")
```

---

### Task 04 📍🔎

<!-- BEGIN QUESTION -->

Why is a permutation test appropriate to use in this situation? In your response, relate this situation to the examples from a class lecture, discussion topic, or the [Chapter 12](https://inferentialthinking.com/chapters/12/1/AB_Testing.html) content.

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

### Task 05 📍

To test their hypotheses, the data scientists have followed our textbook's advice and chosen a test statistic where the following statement is true: Large values of the test statistic favor the alternative hypothesis.

The data scientists' test statistic is one of the two options below. Which one is it? Assign the appropriate number to the variable `correct_test_stat`.

1. "male age average - female age average" in a sample created by randomly shuffling the male/female labels
2. "|male age average - female age average|" in a sample created by randomly shuffling the male/female labels

_Points:_ 3

```python
correct_test_stat = ...
correct_test_stat
```

```python
grader.check("task_05")
```

---

### Task 06 📍

Complete the cell below so that `observed_statistic_ab` evaluates to the observed value of the data scientists' test statistic. Use as many lines of code as you need, and remember that you can use any quantity, table, or array that you created earlier.

_Points:_ 3

```python
observed_statistic_ab = ...
observed_statistic_ab
```

```python
grader.check("task_06")
```

---

### Task 07 📍

Assign `shuffled_labels` to an array of shuffled male/female labels. The rest of the code puts the array in a table along with the data in `sampled_ages`.

*Note:* Check out [12.1](https://inferentialthinking.com/chapters/12/1/AB_Testing.html#predicting-the-statistic-under-the-null-hypothesis) for a refresher on random permutations.

_Points:_ 3

```python
shuffled_labels = ...
original_with_shuffled_labels = sampled_ages.with_columns('Shuffled Label', shuffled_labels)
original_with_shuffled_labels
```

```python
grader.check("task_07")
```

---

### Task 08 📍

The comparison below uses the array `shuffled_labels` from Task 07 and the count `num_females` from Task 01.

For this comparison, assign the correct number from one of the following options to the variable `correct_q8`. **Pretend this is a midterm problem and solve it without doing the calculation in a code cell.**

`comp = np.count_nonzero(shuffled_labels == 'female') == num_females`

1. `comp` is set to `True`.
2. `comp` is set to `False`.
3. `comp` is set to `True` or `False`, depending on how the shuffle came out.

_Points:_ 3

```python
correct_q8 = ...
correct_q8
```

```python
grader.check("task_08")
```

---

### Task 09 📍

Define a function `simulate_one_statistic` that takes no arguments and returns one simulated value of the test statistic. We've given you a skeleton, but feel free to approach this question in a way that makes sense to you. Use as many lines of code as you need. Refer to the code you have previously written in this problem, as you might be able to re-use some of it.

```python
def simulate_one_statistic():
    "Returns one value of our simulated test statistic"
    shuffled_labels = ...
    shuffled_tbl = ...
    group_means = ...
    ...
```

After you have defined your function, run the following cell a few times to see how the statistic varies.

_Points:_ 3

```python
simulate_one_statistic()
```

```python
grader.check("task_09")
```

---

### Task 10 📍

Complete the cell to simulate 4,000 values of the statistic. We have included the code that draws the empirical distribution of the statistic and shows the value of `observed_statistic_ab` from Task 06. Feel free to use as many lines of code as you need.

*Note:* This cell will take around one minute to run.

_Points:_ 3

```python
repetitions = 4000

simulated_statistics_ab = make_array()
...
    simulated_statistics_ab = ...

# Do not change these lines
Table().with_columns('Simulated Statistic', simulated_statistics_ab).hist()
plt.scatter(observed_statistic_ab, -0.002, color='red', s=70);
```

```python
grader.check("task_10")
```

---

### Task 11 📍

Use the simulation to find an empirical approximation to the p-value. Assign `p_val` to the appropriate p-value from this simulation.

_Points:_ 2

```python
p_val = ...
p_val
```

```python
grader.check("task_11")
```

---

### Task 12 📍🔎

<!-- BEGIN QUESTION -->

What conclusion can we draw from this permutation test? In your response, compare the p-value to an assumed significance level of 5% and indicate whether the data and test results support the null or" alternative hypothesis.

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

Follow these steps to submit your homework assignment:

1. **Review the Rubric:** View the rubric on the associated Canvas Assignment page to understand the scoring criteria.
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download IPYNB` to download the notebook (`.ipynb`) file.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.ipynb` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 3)
```
>>> 0 <= num_females <= 500
True
```
failure_message: ❌ num_females should be a value betwen 0 and 500.
success_message: ✅ num_females is a value betwen 0 and 500.


Tests for task_02 (points: 3)
```
>>> import numpy as np
>>> type(avg_male_vs_female) in set([bool, np.bool_])
True
```
failure_message: ❌ avg_male_vs_female should be a bool.
success_message: ✅ avg_male_vs_female is a bool.


Tests for task_03 (points: 3)
```
>>> type(null_statement_number) == int
True
```
failure_message: ❌ null_statement_number should be an int.
success_message: ✅ null_statement_number is an int.
```
>>> type(alternative_statement_number) == int
True
```
failure_message: ❌ alternative_statement_number should be an int.
success_message: ✅ alternative_statement_number is an int.
```
>>> any((null_statement_number == x for x in np.arange(1, 7)))
True
```
failure_message: ❌ null_statement_number should be a number between 1 and 6.
success_message: ✅ null_statement_number is a number between 1 and 6.
```
>>> any((alternative_statement_number == x for x in np.arange(1, 7)))
True
```
failure_message: ❌ alternative_statement_number should be a number between 1 and 6.
success_message: ✅ alternative_statement_number is a number between 1 and 6.
```
>>> null_statement_number != alternative_statement_number
True
```
failure_message: ❌ null_statement_number should not the same as alternative_statement_number.
success_message: ✅ null_statement_number is not the same as alternative_statement_number.


Tests for task_05 (points: 3)
```
>>> type(correct_test_stat) == int
True
```
failure_message: ❌ correct_test_stat should be an int.
success_message: ✅ correct_test_stat is an int.
```
>>> any((correct_test_stat == x for x in np.arange(1, 3)))
True
```
failure_message: ❌ correct_test_stat should be either 1 or 2.
success_message: ✅ correct_test_stat is either 1 or 2.


Tests for task_06 (points: 3)
```
>>> type(observed_statistic_ab) == float
True
```
failure_message: ❌ observed_statistic_ab should be a float.
success_message: ✅ observed_statistic_ab is a float.
```
>>> observed_statistic_ab >= 0
True
```
failure_message: ❌ observed_statistic_ab should be 0 or larger.
success_message: ✅ observed_statistic_ab is 0 or larger.


Tests for task_07 (points: 3)
```
>>> type(original_with_shuffled_labels) == Table
True
```
failure_message: ❌ original_with_shuffled_labels should be a Table.
success_message: ✅ original_with_shuffled_labels is a Table.
```
>>> set(original_with_shuffled_labels.labels) == {'Gender', 'Age', 'Shuffled Label'}
True
```
failure_message: ❌ The labels for original_with_shuffled_labels are not correct.
success_message: ✅ The labels for original_with_shuffled_labels are correct.
```
>>> original_with_shuffled_labels.num_rows == 500
True
```
failure_message: ❌ original_with_shuffled_labels doesn't have the correct number of rows.
success_message: ✅ original_with_shuffled_labels has the correct number of rows.


Tests for task_08 (points: 3)
```
>>> type(correct_q8) == int
True
```
failure_message: ❌ correct_q8 should be an int.
success_message: ✅ correct_q8 is an int.
```
>>> any((correct_q8 == x for x in np.arange(1, 4)))
True
```
failure_message: ❌ correct_q8 should be a number between 1 and 3.
success_message: ✅ correct_q8 is a number between 1 and 3.


Tests for task_09 (points: 3)
```
>>> reps = 100
>>> num_within_5 = sum((-5 < simulate_one_statistic() < 5 for _ in range(reps)))
>>> num_within_5 / reps >= 0.95
True
```
failure_message: ❌ simulate_one_statistic() should return values between -5 and 5.
success_message: ✅ simulate_one_statistic() returns values between -5 and 5.
```
>>> np.random.seed(2024)
>>> np.isclose(simulate_one_statistic(), 0.5769230769230802)
True
```
failure_message: ❌ simulate_one_statistic doesn't seem to be working correctly.
success_message: ✅ simulate_one_statistic seems to be working correctly


Tests for task_10 (points: 3)
```
>>> len(simulated_statistics_ab) == 4000
True
```
failure_message: ❌ simulated_statistics_ab should have 4,000 elements.
success_message: ✅ simulated_statistics_ab has 4,000 elements.


Tests for task_11 (points: 2)
```
>>> import numpy as np
>>> type(p_val) in set([float, np.float32, np.float64])
True
```
failure_message: ❌ p_val should be a float.
success_message: ✅ p_val is a float.
```
>>> 0 <= p_val <= 1
True
```
failure_message: ❌ p_val should be a number between 0 and 1.
success_message: ✅ p_val is a number between 0 and 1.
